<!DOCTYPE HTML>
<html>
    <head>
        <title>CSS Regions: onkeydown events work outside region</title>
        <link rel="author" title="David Alcala" href="mailto:dalcala@adobe.com">
        <link rel="help" href="http://www.w3.org/TR/css3-regions/">
        <link rel="help" href="http://www.w3.org/TR/css3-regions/#the-flow-into-property">
        <link rel="help" href="http://www.w3.org/TR/css3-regions/#the-flow-from-property">
        <link rel="help" href="http://www.w3.org/TR/css3-regions/#relation-to-document-events">
        <meta name="assert" content="CSS regions module does not alter the normal processing of events in the document
         tree. Therefore, onkeydown and onkeyup events inside and outside a named flow should function correctly.">
        <meta name="flags" content="dom interact">
        <link href="support/css/regions-keyboard.css" rel="stylesheet" type="text/css">
        <link href="/resources/testharness.css" rel="stylesheet" type="text/css">
        <script src="/resources/testharness.js"></script>
    </head>
    <body>
        <p id="msg">
            keydown any key in the text field in the orange div, and do not keyup.
        </p>
        <div id="content">
            <div class="block" title="FAIL"></div>
            <div class="block" title="FAIL"></div>
            <div id="target-block" class="block">
                <form id="my-form">
                    <input type="text" id="textfield-inside" size="10">
                    <input type="button" id="myButton" value="Button">
                </form>
            </div>
        </div>
        <div class="region"></div>
        <div class="region"></div>
        <div class="region"></div>
        <br/>
        <div id="outside-region">This div is outside the region
            <form id="form-outside">
                <input type="text" id="textfield-outside" size="10">
                <input type="button" id="myButton" value="Button">
            </form>
        </div>
        <div id="log"></div>
    </body>
    <script type="text/javascript">
     
    //   runMouseTest();
        
        var testTimeout = 10000;
        setup({timeout: testTimeout});
        
        // This block is executed if running in WebKit's harness
        if (window.testRunner)
        {
        	testRunner.dumpAsText(false);
        }

        // Verify that CSS Regions are enabled in the browser.
        // Divs will be horizontal if Regions are enabled.
        // Divs will be vertical if Regions are not enabled.
        function checkLeftPosition(elemID) 
        {
            return document.getElementById(elemID).getBoundingClientRect().left;
        }
        
      //  function runMouseTest()
      //  {
        
        if(checkLeftPosition("target-block") != 0) {
                
            var test1 = async_test("onkeydown event fired in textfield outside region", {timeout: testTimeout});
                test1.step( function() {
                    var testTarget = document.getElementById("outside-region");
                    testTarget.onkeydown = test1.step_func( function(evt) {
                        /* Log assertion that the event got fired */
                        assert_true(true);

                        /* Verify the target outside the region has the correct ID */
                        test(function() { assert_equals(evt.target.id, "textfield-outside")}, "Target outside region has the correct ID");

                        test1.done();
                    });
                    keyDown("textfield-outside");
                });

         //   }  end of runMouseTest()

            function keyDown(block)
            {
            	if(window.testRunner)
               	{
            		var input = document.getElementById(block);
                    input.focus();
                    eventSender.keyDown('a');
                    console.log(input.id);
                    console.log(input.value);
            	}
            }

        } //end of checkLeftPosition test
        else { 
            test(function() {assert_true(false) }, "Regions are not enabled");
        }

    </script>
</html>