<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: The template contents owner document type is HTML document</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="assert" content="The template contents owner document type is HTML document, if template is declared in HTML document">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#definitions">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

// document has browsing context
testInIFrame(null, function(context) {
    var doc = context.iframes[0].contentDocument;
    var template = doc.createElement('template');

    var div = doc.createElement('div');
    div.setAttribute('id', 'div1');
    div.innerText = 'Some text';

    template.appendChild(div);

    doc.body.appendChild(template);

    assert_equals(template.content.ownerDocument.constructor.name, 'HTMLDocument',
            'Template content owner should be a HTML document');

}, 'The template contents owner document type (document has browsing context)');


// document has no browsing context
test(function() {
    var doc = newHTMLDocument();
    var template = doc.createElement('template');

    doc.body.appendChild(template);

    assert_equals(template.content.ownerDocument.constructor.name, 'HTMLDocument',
            'Template content owner should be a HTML document');

}, 'The template contents owner document type (document has no browsing context)');
</script>
</body>
</html>
