<!DOCTYPE html>
<html>
<head>
    <title>CSSOM View - 4.2 - screen.width mutation throws exception</title>
    <link rel="author" title="Neils Christoffersen" href="mailto:neils.christoffersen@gmail.com">
    <link rel="help" href="http://www.w3.org/TR/cssom-view/#the-screen-interface">
    <meta name="flags" content="dom">
    <meta name="assert" content="screen.width mutation throws exception"> 
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script> 
</head>
<body>
    <div id="myDiv"></div>
    <div id="log"></div>
    <script>
		test(function() {
			assert_throws(null, function() {
				window.screen.width = 0 
			}), "chaning window.screen.width should throw exception"
		}, "mutation exception test");		
    </script>    
</body>
</html>
