<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Template element as a descendant of the head element.</title>
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="assert" content="Template element can be a descendant of the head element">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#template-element">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

testInIFrame('../resources/template-descendant-head.html', function(ctx) {
    var doc = ctx.iframes[0].contentDocument;

    assert_true(doc.head.querySelector('template') != null,
        'Template element should be a descendant of the head element');

}, 'Template element as a descendant of the head element. Test loading from a file');


test(function() {
    var doc = newHTMLDocument();

    doc.head.innerHTML = '<template>some text</template>';

    assert_true(doc.head.querySelector('template') != null,
        'Template element should be a descendant of the head element');

}, 'Template element as a descendant of the head element. '
    + 'The template element is assigned to head\'s innerHTML');


test(function() {
    var doc = newHTMLDocument();

    var template = doc.createElement('template');

    doc.head.appendChild(template);

    assert_true(doc.head.querySelector('template') != null,
        'Template element should be a descendant of the head element');

}, 'Template element as a descendant of the head element. '
    + 'Test template element appended to head element by appendChild()');


</script>
</body>
</html>
