<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: HEAD element cannot be content of TEMPLATE element</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="assert" content="HEAD element cannot be content of TEMPLATE element">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#template-element">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">


test(function() {
    var doc = newHTMLDocument();
    var template = doc.createElement('template');

    template.innerHTML = '<head></head>';

    doc.body.appendChild(template);

    assert_equals(template.content.childNodes.length, 0,
            'Template cannot contain HEAD element');

}, 'HEAD element cannot be content of TEMPLATE element. '
    + 'Test empty HEAD element assigned to template innerHTML');


test(function() {
    var doc = newHTMLDocument();
    var template = doc.createElement('template');

    template.innerHTML = '<head><title>test</title></head>';

    doc.body.appendChild(template);

    assert_equals(template.content.childNodes.length, 1,
            'Wrong number of template content children');
    assert_equals(template.content.firstChild.nodeName, 'TITLE',
            'Template should contain children of ignored HEAD element');

}, 'HEAD element cannot be content of TEMPLATE element. '
    + 'Test not empty HEAD element assigned to template innerHTML');


test(function() {
    var doc = newHTMLDocument();
    var template = doc.createElement('template');

    template.innerHTML = '<div id="dv">Some text</div><head><title>test</title></head>';

    doc.body.appendChild(template);

    assert_equals(template.content.childNodes.length, 2,
            'Wrong number of template content children');
    assert_true(template.content.querySelector('#dv') != null,
            'Template should contain valid element');
    assert_equals(template.content.lastChild.tagName, 'TITLE',
            'Template should contain children of ignored HEAD element');

}, 'HEAD element cannot be content of TEMPLATE element. '
    + 'Test HEAD element and some valid element before it, assigned to template innerHTML');


test(function() {
    var doc = newHTMLDocument();
    var template = doc.createElement('template');

    template.innerHTML = '<head><title>test</title></head><div id="dv">Some text</div>';

    doc.body.appendChild(template);

    assert_equals(template.content.childNodes.length, 2,
            'Wrong number of template content children');
    assert_equals(template.content.firstChild.tagName, 'TITLE',
            'Template should contain children of ignored HEAD element');
    assert_true(template.content.querySelector('#dv') != null,
            'Template should contain valid element');

}, 'HEAD element cannot be content of TEMPLATE element. '
    + 'Test HEAD element and some valid element after it, assigned to template innerHTML');


test(function() {
    var doc = newHTMLDocument();
    var template = doc.createElement('template');

    template.innerHTML = '<template id="t2"><head><title>test</title></head></template>';

    doc.body.appendChild(template);

    assert_equals(template.content.childNodes.length, 1,
            'Template should contain nested template');
    assert_true(template.content.querySelector('#t2') != null,
            'Template should contain nested element');

    var nestedTemplate = template.content.querySelector('#t2');

    assert_equals(nestedTemplate.content.childNodes.length, 1,
            'Wrong number of template content children');
    assert_equals(nestedTemplate.content.firstChild.tagName, 'TITLE',
            'Template should contain children of ignored HEAD element');

}, 'HEAD element cannot be content of TEMPLATE element. '
    + 'Test HEAD tag inside template tag assigned to another template\'s innerHTML');


testInIFrame('../resources/template-contents-head.html', function(ctx){
    var doc = ctx.iframes[0].contentDocument;

    var template = doc.body.querySelector('template');

    assert_equals(template.content.childNodes.length, 0,
            'Template cannot contain HEAD element');
}, 'HEAD element cannot be content of TEMPLATE element. '
    + 'Test loading a HTML file with HEAD tag inside template');

</script>
</body>
</html>
