<!DOCTYPE HTML>
<head>
<meta charset="utf-8">
<title>PerformanceObservers: changing observe behavior in callback</title>
<meta name="author" title="JosephPecoraro" href="mailto:joepeck@webkit.org">
<meta name="assert" content="PerformanceObserverEntryList changing observe behavior in callback">
<link rel="help" href="https://w3c.github.io/performance-timeline/#performance-timeline">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src="performanceobservers.js"></script>
</head>
<h1>PerformanceObservers: changing observe behavior in callback</h1>
<div id="log"></div>
<script>
  async_test(function (t) {
    var callbackCount = 0;
    var observer = new PerformanceObserver(
        t.step_func(function (entryList, obs) {
          callbackCount++;

          if (callbackCount === 1) {
            checkEntries(entryList.getEntries(), [
              {entryType: "measure", name: "measure1"},
            ]);
            observer.observe({entryTypes: ["mark"]});
            performance.mark("mark2");
            performance.measure("measure2");
            return;
          }

          if (callbackCount === 2) {
            checkEntries(entryList.getEntries(), [
              {entryType: "mark", name: "mark2"},
            ]);
            performance.mark("mark-before-change-observe-state-to-measure");
            performance.measure("measure-before-change-observe-state-to-measure");
            observer.observe({entryTypes: ["measure"]});
            performance.mark("mark3");
            performance.measure("measure3");
            return;
          }

          if (callbackCount === 3) {
            checkEntries(entryList.getEntries(), [
              {entryType: "measure", name: "measure3"},
              {entryType: "mark", name: "mark-before-change-observe-state-to-measure"},
            ]);
            performance.mark("mark-before-change-observe-state-to-both");
            performance.measure("measure-before-change-observe-state-to-both");
            observer.observe({entryTypes: ["mark", "measure"]});
            performance.mark("mark4");
            performance.measure("measure4");
            return;
          }

          if (callbackCount === 4) {
            checkEntries(entryList.getEntries(), [
              {entryType: "measure", name: "measure-before-change-observe-state-to-both"},
              {entryType: "measure", name: "measure4"},
              {entryType: "mark", name: "mark4"},
            ]);
            performance.mark("mark-before-disconnect");
            performance.measure("measure-before-disconnect");
            observer.disconnect();
            performance.mark("mark-after-disconnect");
            performance.measure("measure-after-disconnect");
            t.done();
            return;
          }

          assert_unreached("The callback must not be invoked after disconnecting");
        })
      );

    observer.observe({entryTypes: ["measure"]});
    performance.mark("mark1");
    performance.measure("measure1");
  }, "PerformanceObserver modifications inside callback should update filtering and not clear buffer");
</script>
