<!DOCTYPE html>
<html>
<head>
<title>When an HTML Document is loaded in a browsing context, a new registry must be created and associated with this document.</title>
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="assert" content="When an HTML Document is loaded in a browsing context, a new registry must be created and associated with this document.">
<link rel="help" href="http://www.w3.org/TR/custom-elements/#creating-and-passing-registries">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

testInIFrame('../resources/blank.html', function(context) {
    var doc = context.iframes[0].contentDocument;
    try {
        doc.register('x-element');
    } catch (e) {
        assert_unreached('Unexpected exception, while registering a valid custom element');
    }
}, 'Document with browsing context must have a new empty registry');


testInIFrame('../resources/blank.html', function(context) {
    var doc = context.iframes[0].contentDocument;
    var doc1 = newHTMLDocument();
    doc1.register('x-element');
    try {
        doc.register('x-element');
    } catch (e) {
        assert_unreached('Unexpected exception while registering a valid custom element ' +
        'in a document which has it\'s own registry');
    }
}, 'Document with browsing context must have a new empty registry. ' +
    'Test elements with the same type name in different documents');

</script>
</body>
</html>
