<!DOCTYPE html>
<html>
<head>
<title>Custom element local name should be the name from custom element definition</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="assert" content="When custom element constructor is called then custom element's name should be set to local name in element definition">
<link rel="help" href="http://www.w3.org/TR/custom-elements/#instantiating-custom-elements">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src="../testcommon.js"></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">
test(function() {
    var doc = newHTMLDocument();
    var GeneratedConstructor = doc.registerElement('x-a');
    var customElement = new GeneratedConstructor();
    assert_equals(customElement.localName, 'x-a',
        'Custom element local name should be the name from custom element definition');
}, 'Custom element local name should be the name from custom element definition');


test(function() {
    var doc = newHTMLDocument();
    HTML5_ELEMENTS.forEach(function(element) {
        var obj = doc.createElement(element);
        var name = 'x-b-' + element;
        var proto = Object.create(obj.constructor.prototype);
        var GeneratedConstructor = doc.registerElement(name, {
            prototype: proto,
            extends: element
        });
        var customElement = new GeneratedConstructor();

        assert_equals(customElement.localName, element,
            'Custom element local name should be the name from custom element definition');
    });
}, 'Custom element local name should be the name from custom element definition. ' +
    'Test constructor of extended HTML element');


testInIFrame('../resources/blank.html', function(doc) {
    var GeneratedConstructor = doc.registerElement('x-c');
    var customElement = new GeneratedConstructor();
    assert_equals(customElement.localName, 'x-c',
        'Custom element local name should be the name from custom element definition');
}, 'Custom element local name should be the name from custom element definition. ' +
    'Document has browsing context');


testInIFrame('../resources/blank.html', function(doc) {
    HTML5_ELEMENTS.forEach(function(element) {
        var obj = doc.createElement(element);
        var name = 'x-d-' + element;
        var proto = Object.create(obj.constructor.prototype);
        var GeneratedConstructor = doc.registerElement(name, {
            prototype: proto,
            extends: element
        });
        var customElement = new GeneratedConstructor();

        assert_equals(customElement.localName, element,
            'Custom element local name should be the name from custom element definition');
    });
}, 'Custom element local name should be the name from custom element definition. ' +
    'Test constructor of extended HTML element. Document has browsing context');
</script>
</body>
</html>
