<!DOCTYPE html>
<html>
<head>
<title>Chars allowed for custom element type </title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="assert" content="The custom element type identifies a custom element interface and is a sequence of characters that must match the NCName production">
<link rel="help" href="https://dvcs.w3.org/hg/webcomponents/raw-file/default/spec/custom/index.html#concepts">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

test(function() {
    var doc = newHTMLDocument();
    for (var i = 0; i < EXTENDER_CHARS.length; i++) {
        var name = 'x-' + String.fromCharCode(EXTENDER_CHARS[i]);
        try {
            doc.register(name);
        } catch (e) {
            assert_unreached('Exception should not be thrown in case of attempt ' +
                    'to register a custom element with a name \'' + name + 
                    '\' (char code U+' + EXTENDER_CHARS[i].toString(16) + ')');
        }        
    }
}, 'Registering valid custom element types with extender char');


test(function() {
    var doc = newHTMLDocument();
    for (var i = 0; i < COMBINING_CHARS.length; i++) {        
        var name = 'x-' + String.fromCharCode(COMBINING_CHARS[i]);
        try {
            doc.register(name);
        } catch (e) {
            assert_unreached('Exception should not be thrown in case of attempt ' +
                    'to register a custom element with a name \'' + name + 
                    '\' (char code U+' + COMBINING_CHARS[i].toString(16) + ')');
        }        
    }
}, 'Registering valid custom element types with combining char');


test(function() {
    var doc = newHTMLDocument();
    for (var i = 0; i < BASE_CHARS_SINGLE.length; i++) {
        var name = 'x-' + String.fromCharCode(BASE_CHARS_SINGLE[i]);
        try {
            doc.register(name);
        } catch (e) {
            assert_unreached('Exception should not be thrown in case of attempt ' +
                    'to register a custom element with a name \'' + name + '\' (char code U+' + 
                    BASE_CHARS_SINGLE[i].toString(16) + ')');
        }
    }
}, 'Registering valid custom element types with single base chars');


test(function() {
    var doc = newHTMLDocument();
    
    // BASE_CHARS_RANGES is a flattened sequence of range start, range end pairs
    for (var i = 0; i < BASE_CHARS_RANGES.length - 1; i += 2) {
        var rangeStart = BASE_CHARS_RANGES[i];
        var rangeEnd = BASE_CHARS_RANGES[i+1];
        
        for (var charCode = rangeStart; charCode <= rangeEnd; charCode++ ) {
            var name = 'x-' + String.fromCharCode(charCode) + '-' + charCode.toString(16);            
            try {
                doc.register(name);
            } catch (e) {
                assert_unreached('Exception should not be thrown in case of attempt ' +
                        'to register a custom element with a name \'' + name + 
                        '\' (char code U+' + charCode.toString(16) + ')');
            }
        }
    }
}, 'Registering valid custom element types with base chars ranges');


test(function() {
    var doc = newHTMLDocument();
    for (var i = 0; i < IDEOGRAPHIC_CHARS_SINGLE.length; i++) {
        var name = 'x-' + String.fromCharCode(IDEOGRAPHIC_CHARS_SINGLE[i]);
        try {
            doc.register(name);
        } catch (e) {
            assert_unreached('Exception should not be thrown in case of attempt ' +
                    'to register a custom element with a name \'' + name + 
                    '\' (char code U+' + IDEOGRAPHIC_CHARS_SINGLE[i].toString(16) + ')');
        }        
    }
}, 'Registering valid custom element types with single ideographic chars');


test(function() {
    var doc = newHTMLDocument();
    
    // IDEOGRAPHIC_CHARS_RANGES is a flattened sequence of range start, range end pairs
    for (var i = 0; i < IDEOGRAPHIC_CHARS_RANGES.length - 1; i += 2) {
        var rangeStart = IDEOGRAPHIC_CHARS_RANGES[i];
        var rangeEnd = IDEOGRAPHIC_CHARS_RANGES[i+1];
        
        for (var charCode = rangeStart; charCode <= rangeEnd; charCode++ ) {
            var name = 'x-' + String.fromCharCode(charCode) + '-' + charCode.toString(16);
            try {
                doc.register(name);
            } catch (e) {
                assert_unreached('Exception should not be thrown in case of attempt ' +
                        'to register a custom element with a name \'' + name + 
                        '\' (char code U+' + charCode.toString(16) + ')');
            }
        }
    }
}, 'Registering valid custom element types with ideographic chars ranges');


test(function() {
    var doc = newHTMLDocument();
    
    // DIGIT_CHARS_RANGES is a flattened sequence of range start, range end pairs
    for (var i = 0; i < DIGIT_CHARS_RANGES.length - 1; i += 2) {
        var rangeStart = DIGIT_CHARS_RANGES[i];
        var rangeEnd = DIGIT_CHARS_RANGES[i+1];
        
        for (var charCode = rangeStart; charCode <= rangeEnd; charCode++ ) {
            var name = 'x-' + String.fromCharCode(charCode) + '-' + charCode.toString(16);
            try {
                doc.register(name);
            } catch (e) {
                assert_unreached('Exception should not be thrown in case of attempt ' +
                        'to register a custom element with a name \'' + name + 
                        '\' (char code U+' + charCode.toString(16) + ')');
            }
        }
    }
}, 'Registering valid custom element types with digit chars');

</script>
</body>
</html>
