<!DOCTYPE html>
<html>
<head>
<title>Chars allowed for custom element type </title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="assert" content="The custom element type identifies a custom element interface and is a sequence of characters that must match the NCName production">
<link rel="help" href="https://dvcs.w3.org/hg/webcomponents/raw-file/default/spec/custom/index.html#concepts">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

function checkValidName(doc, name) {
    try {
        doc.registerElement(name);
    } catch (e) {
        assert_unreached('The custom element name \'' + name +
               '\' should be registered without errors');
    }
}

// check each element of the supplied array with the checkFunction
function testCharCodeArray(doc, charCodeArray, checkFunction) {
    for (var i = 0; i < charCodeArray.length; i++){
        var c = charCodeArray[i];
        var name = 'x-' + String.fromCharCode(c) + '-x' + c.toString(16);
        checkFunction(doc, name);
    }
}

// check each element of the supplied ranges array with the checkFunction
function testCharCodeRangesArray(doc, charCodeRangesArray, checkFunction) {
    for (var i = 0; i < charCodeRangesArray.length; i += 2){
        var rangeStart = charCodeRangesArray[i];
        var rangeEnd = charCodeRangesArray[i+1];
        for (var c = rangeStart; c <= rangeEnd; c++){
            var name = 'x-' + String.fromCharCode(c) + '-x' + c.toString(16);
            checkFunction(doc, name);
        }
    }
}

test(function() {
    testCharCodeArray(newHTMLDocument(), BASE_CHARS_SINGLE, checkValidName);
}, 'Registering valid custom element types with single base chars');

test(function() {
    testCharCodeRangesArray(newHTMLDocument(), BASE_CHARS_RANGES, checkValidName);
}, 'Registering valid custom element types with base chars ranges');

test(function() {
    testCharCodeArray(newHTMLDocument(), IDEOGRAPHIC_CHARS_SINGLE, checkValidName);
}, 'Registering valid custom element types with single ideographic chars');

test(function() {
    testCharCodeRangesArray(newHTMLDocument(), IDEOGRAPHIC_CHARS_RANGES, checkValidName);
}, 'Registering valid custom element types with ideographic chars ranges');

test(function() {
    testCharCodeRangesArray(newHTMLDocument(), DIGIT_CHARS_RANGES, checkValidName);
}, 'Registering valid custom element types with digit chars');

test(function() {
    testCharCodeArray(newHTMLDocument(),
                      ['.'.charCodeAt(0), '-'.charCodeAt(0), '_'.charCodeAt(0)],
                      checkValidName);
}, 'Registering valid custom element types with characters \'.\', \'-\', \'_\'');

test(function() {
    testCharCodeArray(newHTMLDocument(), COMBINING_CHARS, checkValidName);
}, 'Registering valid custom element types with combining char');

test(function() {
    testCharCodeArray(newHTMLDocument(), EXTENDER_CHARS, checkValidName);
}, 'Registering valid custom element types with extender char');

</script>
</body>
</html>
